# Homework 05

In this exercise, you will practice inferential statistics with confidence intervals, bootstrapping, and hypothesis testing. Problems may involve a combination of math and code. 

Recall that you can use LaTeX to nicely format your math inside Markdown cellsby enclosing equations in single dollar signs (e.g., $x^2+4=8$) for inline math or double dollar signs for centered equations like $$P(X > 5) = \frac{1}{6}.$$ For a reference if you are new to LaTeX, see the [overleaf documentation for mathematical expressions](https://www.overleaf.com/learn/latex/mathematical_expressions). **Please note that Gradescope currently does not render ```\begin{align} ... \end{align}``` and ```\begin{equation} ... \end{equation}```. Consider using ```\begin{aligned} ... \end{aligned}``` as an alternative to save the graders some hassle.**

Show your work and/or briefly explain your answers. In general, you will not receive full credit for numeric answers with no accompanying work or justification (math, code, explanation). For numeric answers, we will accept answers that are very slightly off due to rounding, z score of 2 vs. 1.96, etc. 

When you finish please go to Kernel --> Restart and Run All, and then double check that your notebook looks correct before submitting your .ipynb file (the notebook file) on gradescope.

```python
# Run this code cell to import relevant libraries
import numpy as np
import pandas as pd
from scipy import stats
```

### Question 1

1. A website is trying to increase registration for first-time visitors, exposing a random subset of these visitors to a new site design. Of $752$ randomly sampled visitors over a month who saw the new design, $64$ registered. Construct a $95\%$ confidence interval for the percentage of visitors who would register for the website under the new design using the normal distribution. Save your answer in a tuple `q1_1` with two `numpy.float64` items that `q1_1[0]` is the left bound and `q1_1[1]` is the right bound. Use percentages for both bounds, for example, save $50.0$ in your answer for $50\%$ of visitors.
2. A study examined the average pay for a random sample of men and women entering the workforce as doctors for $21$ different positions. If each gender was equally paid, then we would expect about half of those positions to have men paid more than women and women would be paid more than men in the other half of positions. In the study, men were, on average, paid more in $17$ of the $21$ positions. Complete a hypothesis test using the normal distribution (two-sided or one-sided, just be clear which you are reporting) to examine whether there is significant evidence (at the $0.05$ level) of gender discrimination in pay in these positions. Report your p-value and interpret the result. Save your p-value in `q1_2` as a `numpy.float64` and interpret it in the **Answer 1.2** cell.

```python
# Code for question 1
# Leave these lines here for grading and ease of debugging
print(q1_1)

# Leave these lines here for grading and ease of debugging
print(q1_2)
```

### Answer 1.2
Write your interpretation for 1.2 here. There is no need to write anything for 1.1.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Movie Ratings Data
In the remainder of this assignment you will work with the movielens dataset of movie ratings that we have seen before. Below we import and preview the data. It consists of 2 tables: `users` has a row for every individual who has rated any movies, `movie-ratings` has a row for every rating of a particular movie by a particular user. This means users with multiple ratings are in the `movie_ratings` multiple times. The data is a random sample of all of the movie ratings made on the movielens service.

```python
users = pd.read_csv("users.csv")
users.head()
```

```python
movie_ratings = pd.read_csv("movies-all.csv")
movie_ratings.head()
```

### Question 2
1. Compute a $95\%$ confidence interval for the mean `age` of users using the normal distribution. Save your answer in a tuple `q2_1` with two `numpy.float64` items that `q2_1[0]` is the left bound and `q2_1[1]` is the right bound.
2. Compute a $95\%$ confidence interval for the mean `age` of users who have rated the movie `Casablanca (1942)` using the normal distribution. Save your answer in `q2_2`, similar requirements as above.
3. *Casablanca* is an old movie, one might suspect that it has been rated by older individuals on average than the entire dataset. Just looking at the confidence intervals you computed in steps 1 and 2, can you conclude that there is significant evidence for this belief? Why or why not? Put your answer in the **Answer 2.3** cell.

```python
# Code for question 2
# Leave these lines here for grading and ease of debugging
print(q2_1)

# Leave these lines here for grading and ease of debugging
print(q2_2)
```

### Answer 2.3
Write your interpretation for 2.3 here. There is no need to write anything for 2.1 & 2.2.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 3
Only $18$ users have rated the movie `Lost in Space (1998)`.
1. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the average `age` of users who have rated `Lost in Space (1998)`. Save your answer in a tuple `q3_1` with two `numpy.float64` items that `q3_1[0]` is the left bound and `q3_1[1]` is the right bound.
2. One of the advantages of bootstrapping is that we can easily compute confidence intervals for arbitrary measurements of distributions. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the **median** `rating` of `Lost in Space (1998)`. Note that numpy provides a vectorized function for [calculating the median](https://numpy.org/doc/stable/reference/generated/numpy.median.html) as well as the mean. Save your answer in a tuple `q3_2` with two `numpy.float64` items. Similar requirements as above.

```python
# Code for question 3
# Leave these lines here for grading and ease of debugging
print(q3_1)

# Leave these lines here for grading and ease of debugging
print(q3_2)
```

### Question 4
The `Star Wars (1977)` film is quite popular, with a median rating of `5` out of $5$. However, of those that left a rating, male users gave it a slightly higher average rating of about $4.4$ whereas female users gave the same movie an average rating of about $4.2$.

1. Consider the null hypothesis that the average rating of `Star Wars (1977)` is the same for `sex='F'` and `sex='M'` users. The alternative hypothesis is that the average ratings are not equal. Conduct a two-sided t test using [`stats.ttest_ind`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind.html) to evaluate this using the sample ratings data. Report your p-value and interpret it at a significance level of $0.05$. Save your p-value in `q4_1` as a `numpy.float64` and interpret it in the **Answer 4** cell.

2. Consider the null hypothesis that $51\%$ of men would rate `Star Wars (1977)` a `5`. Conduct a two-sided hypothesis test using [`stats.t.cdf`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.t.html) to evaluate this in light of the sample ratings data of male users who rated `Star Wars (1977)`. Report your p-value and interpret it at a significance level of $0.05$. Save your p-value in `q4_2` as a `numpy.float64` and interpret it in the **Answer 4** cell.

3. Consider the null hypothesis that women and men were equally likely to rate `Star Wars (1977)` a `5`. Conduct a two-sided t test using [`stats.ttest_ind_from_stats`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind_from_stats.html) to evaluate this in light of the sample data of female and male users who rated `Star Wars (1977)`. Report your p-value and interpret it at a significance level of $0.05$. Save your p-value in `q4_3` as a `numpy.float64` and interpret it in the **Answer 4** cell. You should observe a different p-value than in step 2 despite the hypotheses under consideration being ostensibly similar. Briefly explain why you observe this difference.

```python
# Code for question 4
# Leave these lines here for grading and ease of debugging
print(q4_1)

# Leave these lines here for grading and ease of debugging
print(q4_2)

# Leave these lines here for grading and ease of debugging
print(q4_3)
```

### Answer 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. We recommend you go to the Kernel option and choose \"Restart & Run All.\" Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/).

Tests for q1 (points: [1, 1, 2, 1, 1])
```
>>> isinstance(q1_1, tuple)
True
```
```
>>> # Looks like you might not have multiplied your confidence interval by 100.
>>> [v >= 1 and v <= 100 for v in q1_1]
[True, True]
```
```
>>> from auto_checks import check_q1_1
>>> check_q1_1(q1_1)
'PASS'
```
```
>>> import numpy
>>> isinstance(q1_2, numpy.float64) or isinstance(q1_2, float)
True
```
```
>>> from auto_checks import check_q1_2
>>> check_q1_2(q1_2)
'PASS'
```


Tests for q2 (points: [1, 1, 2, 1, 1, 2])
```
>>> isinstance(q2_1, tuple)
True
```
```
>>> [v >= 1 and v <= 100 for v in q2_1]
[True, True]
```
```
>>> from auto_checks import check_q2_1
>>> check_q2_1(q2_1)
'PASS'
```
```
>>> isinstance(q2_2, tuple)
True
```
```
>>> [v >= 1 and v <= 100 for v in q2_2]
[True, True]
```
```
>>> from auto_checks import check_q2_2
>>> check_q2_2(q2_2)
'PASS'
```


Tests for q3 (points: [1, 1, 2, 1, 1, 2])
```
>>> isinstance(q3_1, tuple)
True
```
```
>>> [v >= 1 and v <= 100 for v in q3_1]
[True, True]
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> isinstance(q3_2, tuple)
True
```
```
>>> [v >= 0 and v <= 5 for v in q3_2]
[True, True]
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 1, 1, 1])
```
>>> import numpy
>>> isinstance(q4_1, numpy.float64) or isinstance(q4_1, float)
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> import numpy
>>> isinstance(q4_2, numpy.float64) or isinstance(q4_2, float)
True
```
```
>>> from auto_checks import check_q4_2
>>> check_q4_2(q4_2)
'PASS'
```
```
>>> import numpy
>>> isinstance(q4_3, numpy.float64) or isinstance(q4_3, float)
True
```
```
>>> from auto_checks import check_q4_3
>>> check_q4_3(q4_3)
'PASS'
```
