# Homework 06
In this assignment, you will demonstrate your mastery of aggregating and summarizing data with Pandas `groupby` and `pivot_table` and merging/joining datasets using Pandas `concat` and `merge`.

## Part 1: COVID-19 Dataset
In this part we work with the `covid_19_data.csv` dataset. It contains records of the cumulative numbers of confirmed cases of the disease (the `Confirmed` column), deaths from the disease (the `Deaths` column), and recovered individuals (the `Recovered` column). **Remember** that *cumulative* means total up to this date, i.e., a given observation is the previous observation plus any additional since then.

Each such value corresponds to a particular `Province/State` within a given `Country/Region`. Some rows do not include information at the `Province/State` level, in which case the data correspond to the `Country/Region` as a whole. The date of the record is given in the `ObservationDate` column. Below we import the dataset and preview the first few rows. 

```python
# Run but do not modify this code
import pandas as pd
covid = pd.read_csv("covid-19.csv", parse_dates=["ObservationDate"])
covid.head()
```

```python
type(covid["ObservationDate"][0])
```

### Question 1 (12 points, 4 each)

1. Which `Country/Region` has reported the most total observations (rows), and how many did they report? Save the name of the Contry/Region in `q1_1_1` as `str`, and save the number they report in `q1_1_2` as `int` or `np.int64`.
2. How many `Country/Region`s have reported at least 100 `Deaths` due to COVID-19? Save your result in `q1_2`.
3. Which five `Province/State`s of the `Country/Region` of `Mainland China` have reported the most `Deaths` to date, and how many `Deaths` have they reported? Put your answer in `q1_3` such that it is of type `Series` where the index is the `Province/State` and the value is the number of `Deaths` as `float` or `np.float64`.

```python
# Put your code to answer the question here
# Feel free to add additional cells
q1_1_1 = ...
q1_1_2 = ...
q1_2 = ...
q1_3 = ...
```

### Question 2 (8 points, 4 each)

1. Report the first five `Province/State`s within the `Country/Region` of `US` to record a `Confirmed` case along with the dates on which they recorded their first case(s). Put your answer in `q2_1` such that it is of type `Series` where the index is the `Province/State` and the value is the dates on which they recorded their first case(s) as `pandas._libs.tslibs.timestamps.Timestamp`.

2. All of the data for the `Country/Region` of `US` is recorded at the `Province/State` level. Create a Pandas Series called `q2_2` where the primary index has `ObservationDate`s and the values are the total `Confirmed` cases added up across all of the `Province/State`s in the `Country/Region` of `US`. For example, the first several rows of the Series should look like:
```
2020-01-22           1.0
2020-01-23           1.0
2020-01-24           2.0
2020-01-25           2.0
2020-01-26           5.0
```
When you are finished, print `q2_2.head()` and then run the following code cell to visualize the spread of the virus in the United States as a line plot.  

```python
# Put your code to answer the question here
# Feel free to add additional cells
q2_1 = ...

q2_2 = ...
print("2. ", q2_2.head())
```

```python
# Run but do not modify this code
# You should compute `q2_2` above for Q2.2
q2_2.plot()
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 2: Movie Ratings
In this part we return to the movie review dataset that contains information about various users who have rated several different movies. To start, use `first_hundred.csv` contains all data for the first 100 `user_id`s (`1` through `100`) with all information merged into one table. Later, we will consider combining information from additional files with data on the subsequent `user_id`s. Below we import `first_hundred.csv` into the DataFrame `df_fh` and preview the first few rows.

```python
# Run but do not modify this code
import pandas as pd
df_fh = pd.read_csv("first_hundred.csv")
df_fh.head()
```

### Question 3 (8 points, 4 each)
Use Pandas `pivot_table`s to answer the following about the data in `first_hundred.csv`.

1. For each of the five possible `rating` values (1 through 5) report how many ratings of that value were recorded for `sex=='M'` and for `sex=='F'`. That is, how many total ratings of `1` were given by women, how many total ratings of `2` were given by women, etc., and the same for men. Save your result in pivot table `q3_1`, the columns are `sex` and the rows are `rating`.

2. Compute the `decade` of a user to be their `age // 10` (where `//` stands for integer division in Python so that someone with `age` between 0 and 9 would have `decade` 0, between 10 and 19 would have `decade` 1, and so on). Create a pivot table that shows the total number of `rating`s in the dataset for each combination of `decade` and `sex`. That is, your pivot table should show how many total `rating`s were made by users who are `decade==0` and `sex=='F'`, `decade==1` and `sex=='F'`, etc., and the same for `sex=='M'`. Use the `fill_value=0` optional parameter to the Pandas `pivot_table` method so that combinations for which there are no ratings will be printed as 0.  Display your resulting pivot table. Put your result in pivot table `q3_2`, the columns are `sex` and the rows are `decade`.

```python
# Put your code to answer the question here
# Feel free to add additional cells
q3_1 = ...
q3_2 = ...
```

### Question 4 (13 points, part 0 is 1 point, the rest are 4 each)
`first_hundred.csv` only contains ratings from the first 100 `user_id`s containing about 11% of the total ratings. The remaining ratings are split across two additional files: `other_ratings_a.csv` and `other_ratings_b.csv`. In addition, information about additional users not in the first hundred is contained in the `other_users.csv` file, and information about all movies is contained in the `movies.csv` file. You will need to import and appropriately combine data from all of these datasets in order to answer the following questions. We encourage you to start by exploring the individual datasets.

0. To help you check that you combined things correctly, put in `q4_0` the number of rows in your dataframe.

1. What fraction of the ratings were made by women (i.e., `sex=='F'`)? Put your result in `q4_1` as a `float` or `np.float64`.
2. What are the `movie_title`s of the five movies that have received the most number of ratings? Put your answer in `q4_2` such that it is of type `Series` where the index is the `movie_title` and the value is the `rating` as `int` or `np.int64`.
3. Several movies have not received very many `rating`s; average ratings of such movies are not very informative. Instead, what are the `movie_title`s and average `rating`s of the movies that have received *at least 10* `rating`s? Show the five movies with the highest average `rating`. Put your answer in `q4_3` such that it is of type `Series` where the index is the `movie_title` and the value is the average rating as `float` or `np.float64`.

```python
# Put your code to answer the question here
# Feel free to add additional cells

q4_0 = ...

q4_1 = ...

q4_2 = ...

q4_3 = ...
```

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. We recommend you go to the Kernel option and choose "Restart & Run All." Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/)."

Tests for q1 (points: [1, 1, 1, 1, 2, 2, 1, 1, 1, 1])
```
>>> isinstance(q1_1_1, str)
True
```
```
>>> import numpy
>>> isinstance(q1_1_2, int) or isinstance(q1_1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_1_1
>>> check_q1_1_1(q1_1_1)
'PASS'
```
```
>>> from auto_checks import check_q1_1_2
>>> check_q1_1_2(q1_1_2)
'PASS'
```
```
>>> import numpy
>>> isinstance(q1_2, int) or isinstance(q1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_2
>>> check_q1_2(q1_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q1_3, pandas.core.series.Series)
True
```
```
>>> q1_3.shape
(5,)
```
```
>>> from auto_checks import check_q1_3_province
>>> check_q1_3_province(q1_3)
'PASS'
```
```
>>> from auto_checks import check_q1_3_death
>>> check_q1_3_death(q1_3)
'PASS'
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q2_1, pandas.core.series.Series)
True
```
```
>>> q2_1.shape
(5,)
```
```
>>> from auto_checks import check_q2_1_state
>>> check_q2_1_state(q2_1)
'PASS'
```
```
>>> from auto_checks import check_q2_1_date
>>> check_q2_1_date(q2_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q2_2, pandas.core.series.Series)
True
```
```
>>> q2_2.shape
(364,)
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q3_1, pandas.core.frame.DataFrame)
True
```
```
>>> q3_1.shape
(5, 2)
```
```
>>> print(q3_1.index.name, sorted(q3_1.columns))
rating ['F', 'M']
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q3_2, pandas.core.frame.DataFrame)
True
```
```
>>> q3_2.shape
(7, 2)
```
```
>>> print(q3_2.index.name, sorted(q3_2.columns))
decade ['F', 'M']
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 2, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> q4_0
100000
```
```
>>> import numpy
>>> isinstance(q4_1, float) or isinstance(q4_1, numpy.float64)
True
```
```
>>> 0 <= q4_1 <= 1
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_2, pandas.core.series.Series)
True
```
```
>>> q4_2.shape
(5,)
```
```
>>> from auto_checks import check_q4_2_index
>>> check_q4_2_index(q4_2)
'PASS'
```
```
>>> from auto_checks import check_q4_2_rating
>>> check_q4_2_rating(q4_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_3, pandas.core.series.Series)
True
```
```
>>> q4_3.shape
(5,)
```
```
>>> from auto_checks import check_q4_3_index
>>> check_q4_3_index(q4_3)
'PASS'
```
```
>>> from auto_checks import check_q4_3_rating
>>> check_q4_3_rating(q4_3)
'PASS'
```
