# Homework 07
In this assignment, you will practice writing SQL queries to pull data from a relational database. 

## The congress database
A visual schema of the congress database is also available as a pdf along with this file - we encourage you to start by taking a look to get familiar. This database contains the history of the members of the United States congress through the 115th congress (the data end by 2019) as well as a good deal of voting data from 2015-2016. The visual schema shows each table in the database as a yellow box with the table name at the top and the column names listed below. The arrows showing which keys/identifiers match between different tables for join operations. At the bottom you can see previews of the tables. One in particular, the `cur_members` table, contains data about members of the 115th congress (note - everywhere you see us refer to "current" members below, we are referring to those in the `cur_members` table). Many, though not all, of these members are still serving as of the 117th congress, which began in January 2021.

## How to use SQL and report your results

You should use Python sqlite3 and Pandas to run your SQL queries in this assignment. You should include your code (including the Strings containing your SQL queries) in this notebook and show your results for the manual grading. We recommend using multiline strings for your queries for readability; remember that multiline strings in Python are enclosed by triple quotes for example,

```
my_multiline_string = """Hi 
There!"""
```

See the worked examples for an example query and it's result.

## Questions

```python
# Run but do not modify the following code
# to import sqlite3 and pandas, and to connect
# to the congress database
import sqlite3
import pandas as pd
import numpy as np
conn = sqlite3.connect("congress")
# If you cannot load the database and see error message like "no such table", try use absolute path like following:
# import sqlite3
# import pandas as pd
# import numpy as np
# import os.path
# BASE_DIR = os.path.dirname(os.path.abspath("congress"))
# db_path = os.path.join(BASE_DIR, "congress")
# conn = sqlite3.connect(db_path)
```

### Question 1 (4 points)
Who are all of the members of congress from North Carolina (`NC`) in the `cur_members` table? Show all of the information about them from the `cur_members` table. Put your resulting DataFrame in `q1`.

```python
# Start with a SQL query to extract the information
query1 = """ ... """

# Execute queries using pandas
q1 = ...


# Leave this line here so that the table is shown
q1
```

### Question 2 (4 points)
List the five youngest female (`gender = 'F'`) members of the Congress from the `cur_members` table in descending order of `birthday`. Show all information about them from the `cur_members` table. Put your resulting DataFrame in `q2`.

```python
# Start with a SQL query to extract the information
query2 = """ ... """

# Execute queries using pandas
q2 = ...


# Leave this line here so that the table is shown
q2
```

### Question 3 (4 points)
For the current members of congress, show the average, minimum, and maximum age (in years) of the representatives (i.e., `type = 'rep'`). Show the same for senators (i.e., `type = 'sen'`). Compute ages in years as of Feb. 22, 2021 (`'2021-02-22'`). You may compute the age of an individual as simply `'2021-02-22' - birthday`.

Note that this is not *exactly* correct due to date rounding error. A more precise calculation of age in years using [sqlite date and time functions](https://sqlite.org/lang_datefunc.html) looks like `(julianday('2021-02-22')-julianday(birthday)) / 365.25`. We will accept solutions using either, you will get very similar results.

Put your resulting DataFrame in `q3` with colums of `type`, `avg`, `min`, and `max`.

```python
# Start with a SQL query to extract the information
query3 = """ ... """

# Execute queries using pandas & check for the empty string
q3 = ...



# Leave this line here so that the table is shown
q3
```

### Question 4 (4 points)
Which parties (`party`) have had at least 50 senators (`sen`) past and present? For each, show the `party` and the number of senators (`sen`) they have had. 

Hint: Be sure not to double count persons who served multiple terms in the senate. You can use `COUNT(DISTINCT column)` to get the number of distinct values on given `column`.

Put your resulting DataFrame in `q4` with colums of `party` and `num_senators`.

```python
# Start with a SQL query to extract the information
query4 = """ ... """

# Execute queries using pandas
q4 = ...


# Leave this line here so that the table is shown
q4
```

### Question 5 (4 points)
Among the past and present members members of congress who cast at least one vote in the `person_votes` table (i.e., the members whose `person_id`s appear at all in the `person_votes` table), which five members cast the fewest votes? Show their `first_name`s, `last_name`s, and the number of votes they cast in the `person_votes` table. Put your resulting DataFrame in `q5` with colums of `first_name`, `last_name`, and `num_votes`.

Note: There is a vote called "Not Voting", you do not need to treat them specially. Or in other words, do not filter them out when counting, count them as a vote.

```python
# Start with a SQL query to extract the information
query5 = """ ... """

# Execute queries using pandas
q5 = ...


# Leave this line here so that the table is shown
q5
```

### Question 6 (4 points)
In this question, we would like to consider the change in the gender representation of the US House of Representatives over time. Find all of the `start_date`s from the `person_roles` table in which at least 40 women began terms as representatives in the house (i.e., `type='rep'`). Show the `start_date`s and the number of women beginning terms on those dates. Order the results by `start_date` from oldest to most recent. Put your resulting DataFrame in `q6` with colums of `start_date` and `house_women`.

```python
# Start with a SQL query to extract the information
query6 = """ ... """

# Execute queries using pandas
q6 = ...


# Leave this line here so that the table is shown
q6
```

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. We recommend you go to the Kernel option and choose \"Restart & Run All.\" Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/).

Tests for q1 (points: [1, 1, 1, 1])
```
>>> q1.shape
(15, 9)
```
```
>>> print(sorted(q1.columns))
['birthday', 'first_name', 'gender', 'id', 'last_name', 'party', 'religion', 'state', 'type']
```
```
>>> from auto_checks import check_q1_state
>>> check_q1_state(q1)
'PASS'
```
```
>>> from auto_checks import check_q1
>>> check_q1(q1)
'PASS'
```


Tests for q2 (points: [1, 1, 1, 1])
```
>>> q2.shape
(5, 9)
```
```
>>> print(sorted(q2.columns))
['birthday', 'first_name', 'gender', 'id', 'last_name', 'party', 'religion', 'state', 'type']
```
```
>>> from auto_checks import check_q2_gender
>>> check_q2_gender(q2)
'PASS'
```
```
>>> from auto_checks import check_q2
>>> check_q2(q2)
'PASS'
```


Tests for q3 (points: [1, 1, 0.5, 0.5, 0.5, 0.5])
```
>>> q3.shape
(2, 4)
```
```
>>> print(sorted(q3.columns))
['avg', 'max', 'min', 'type']
```
```
>>> from auto_checks import check_q3_type
>>> check_q3_type(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_avg
>>> check_q3_avg(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_min
>>> check_q3_min(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_max
>>> check_q3_max(q3)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 1])
```
>>> q4.shape
(5, 2)
```
```
>>> print(sorted(q4.columns))
['num_senators', 'party']
```
```
>>> from auto_checks import check_q4_party
>>> check_q4_party(q4)
'PASS'
```
```
>>> from auto_checks import check_q4
>>> check_q4(q4)
'PASS'
```


Tests for q5 (points: [1, 1, 2])
```
>>> q5.shape
(5, 3)
```
```
>>> print(sorted(q5.columns))
['first_name', 'last_name', 'num_votes']
```
```
>>> from auto_checks import check_q5
>>> check_q5(q5)
'PASS'
```


Tests for q6 (points: [1, 1, 2])
```
>>> q6.shape
(13, 2)
```
```
>>> print(sorted(q6.columns))
['house_women', 'start_date']
```
```
>>> from auto_checks import check_q6
>>> check_q6(q6)
'PASS'
```
